Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

For this problem set, we'll be using the Jupyter notebook: [![Open In Colab](https://colab.research.google.com/assets/colab-badge.svg)](https://colab.research.google.com/github/GabbySuwichaya/Estimation-Theory-EE523/blob/master/Coding/Step0_Testing.ipynb)

![](jupyter.png)

Please run the following pip command to install the required dependency

```python
!pip install numpy pandas tqdm matplotlib joypy
```

---
## Intro to Jupyter Grading by NB Grader

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Testing 

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
squares(10)
"""Check that squares returns the correct output for several inputs"""
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

---
### Using Latex

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: latex, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Let's get start ... 

 

---

## Estimation Theory (coding examples)

### 1. Convergences in Distribution

#### Definition

If random variables $X_1, X_2, X_3, ...$ **converge in distribution**  to a random variable $X$, shown by $X_n \xrightarrow{d} X$, 



##### Q1. what can we say about the CDF  of X_n and X as $n \rightarrow \infty$ ? 


**Type your latex answer in below markdown**

<!-- grade_id: def, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---

**Central Limit Theorem:** 
One of the applications of Convergences in Distribution is the Central Limit Theorem... 



##### Q2. What is the meaning of CLT? 

*Type your answer in below markdown*

<!-- grade_id: clt, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---

#### Being Simulation for CLT

**Goal:** show that the normalized sum of random variable will converge to a normal distribution.  


**Tasks.**

- Q3. Complete the function `scale_and_normalize_XN()` and `generate_gauss()` that will be used to normalize the sum of Bernoulli sample. 


- Q4. Then, complete the histrogram plotting in `plotting_hist()`


```python
# Installing packages

import numpy as np
import pandas as pd
import scipy.stats as sp
import matplotlib.pyplot as plt
from matplotlib import cm
import math
import os
import joypy
 
cwd = os.getcwd()  
```

##### Q3. Complete the function `scale_and_normalize_XN()` and `generate_gauss()` that will be used to normalize the sum of Bernoulli sample. 

<!-- grade_id: scale_and_normalized, points: 0, grade: false, solution: true -->
```python
def scale_and_normalize_XN(XN, N_round, MU_Xi, STD_Xi):  
    
    mu = N_round*MU_Xi  
    sigma =  np.sqrt(N_round)*STD_Xi # Hint ZN = (XN - mu) / sigma
    
    # YOUR CODE HERE
    raise NotImplementedError()
    return ZN

```

<!-- grade_id: generate_gaussian, points: 0, grade: false, solution: true -->
```python
def generate_gauss(N_size):  
    # Generate Gauss Samples with zero mean and unit variance
    nsigma = 1
    nmu = 0 # Gauss = np.random.standard_normal(N_size)

    
    # YOUR CODE HERE
    raise NotImplementedError()
    return Gauss, nsigma, nmu
```

Begin simulation

<!-- grade_id: pass_simulation, points: 5, grade: true, solution: false -->
```python
N_round = 50 # the number of sum << Vary this one 
N_size  = 1000

########################## The Central Limit Theorem ##########################
# Show that: the sum of RVs will eventually converges to the normal distribution, 
# as the number of sum is getting large.   

orginal_distribution = "Bernoulli"  
p  = 0.5
MU_Xi  = p
STD_Xi = np.sqrt(p*(1-p))
XN = 0
for count_round in range(N_round):
    XN += np.random.binomial(size=N_size, p=p, n=1) 
 
 
# Calculated true mean 
ZN = scale_and_normalize_XN(XN, N_round, MU_Xi, STD_Xi)
 
# Gaussian samples
Gauss, nsigma, nmu = generate_gauss(N_size)

```

##### Q4. Then, complete the histrogram plotting in `plotting_hist()`

<!-- grade_id: plotting_hist, points: 0, grade: false, solution: true -->
```python
def plotting_hist(Gauss, ZN):
    plt.figure(figsize=(10,7)) 
    # YOUR CODE HERE
    raise NotImplementedError()
    plt.ylabel("$P_{X}(x)$")
    plt.xlabel("$x$")
    plt.grid()
    plt.title("Histrogram Comparison of the Sum of Bernoulli. vs. Gaussian Dist.")
    plt.legend()
    plt.savefig("Q1_CompareHistogram.png")
    plt.show()
```

<!-- grade_id: can_it_plt, points: 10, grade: true, solution: false -->
```python
plotting_hist(Gauss, ZN)
```