Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
STUDENT_ID = ""
```

---

# Homework 4

---

All the implementation (36 scores) shall be done in a single Jupyter notebook. 


*Your scores will be given if and only if the `.ipynb` can run successfully and all the  **plots and output showing without any error**. So please read how to submit your code in Readme.md*

## Probabilistic models

---

Suppose that we want to estimate $\bm{X} \in \mathcal{R^N}$ from the observed data $\bm{Y} \in \mathcal{R^N}$ where $\bm{Y} = \bm{X} + N$; where $N\in \mathcal{R}$ is the noise corrupting the target signal. 

The probabilistic relationship between $\bm{X}$ and $\bm{Y}$ can be represented by the following probabilistic models:

- The likelihood is $ f_{\bm{Y}|\bm{X}}(\bm{y}|\bm{x}) = \mathcal{N}(\bm{y}| \bm{x}, \sigma_V^2) \propto \prod_i \frac{1}{2 \pi \sigma_V^2} \exp{ - \frac{(y_i - x_i)^2}{2 \sigma_V^2} }$. 

- The prior for $x$ is a Gaussian-Gamma distribution, that is,  

	- $\bm{X} \sim f_{\bm{X}}(\bm{x})$  is $ f_{X }(\bm{x}| \bm{\mu}, \bm{\tau}) = \mathcal{N}(\bm{x}| \bm{\mu}, \bm{\tau}) \propto  \prod_i \frac{\tau_i}{2 \pi} \exp{ \frac{\tau_i (x_i - \mu_i)^2}{2} }$ 
	- $\bm{\tau} \sim f_{\bm{\tau}}(\bm{\tau}) = \text{Gamma}(\bm{\tau}; \alpha, \beta) \propto \prod_i  \tau_i^{\alpha - 1} \exp{- \beta \tau_i}$  

It is assume that the true distribution of $\bm{X}$ is the same as the prior distribution.  

Therefore, the joint distribution among the random variables are 

$$p(\bm{y}, \bm{x}, \bm{\tau}; \bm{\mu}, \sigma_V^2, \alpha, \beta) = \mathcal{N}(\bm{y}| \bm{x}, \sigma_V^2) \mathcal{N}(\bm{x}| \bm{\mu}, \bm{\tau}) \text{Gamma}(\bm{\tau}; \alpha, \beta).$$

Your observations are the random sample vector $\bm{Y}^1, \bm{Y}^2, ..., \bm{Y}^k$ that are independently and identically distributes according to each element in the multi-variate $\bm{Y}^j$ in $\mathcal{R^N}$.

*If you somehow find out that your answers in the Quiz 2 was not right,  feel free to **correct** them in this one.*

---

```python
! pip install numpy pandas tqdm matplotlib
```

```python
import numpy as np  
import joypy
from matplotlib import cm
import matplotlib.pyplot as plt
import pandas as pd
import os
```

## Estimator Simulation

- [Q 1. Data geneation (6 scores)](#q1-data-generation-total-6-scores)

    -  [Q 1.1 Parameter settings](#q11-parameter-settings)

    -  [Q 1.2 Data generation (2 score)](#q12-data-generation-according-to-the-above-specifications-ie-the-likelihood-and-priors--2-scores)

    -  [Q 1.3 Data visualization (4 score)](#q13-show-that-the-distributions-are-correct-by-histograms-or-any-suitable-visualization-4-scores)


- [Q 2. Implemetation of MLE/MAP (6 scores)](#q2-implementation-of-mapmle-estimator-12-scores)

    - [Q 2.1 Analytical solutions for MAP/MLE estimators (2 scores)](#q21-analytical-solutions-of-mapmle-estimator-for-4-scores)

    - [Q 2.2 Implementation of MAP/MLE estimators (4 scores)](#q-22-implementation-of-mapmle-estimator-for-and-8-scores)


- [Q 3. Verify your results (8 scores)](#q3-verify-your-results-8-scores)

    - [Q 3.1 Check your result with respect to the generated data (4 scores)](#q31-checking-with-respect-to-the-generated-4-scores)

    - [Q 3.2 Verify with KL-Divergence (4 scores)](#q32-verify-with--divergence--4-scores)

 
- [Q 4. FitDist an alternative tool (10 scores)](#q4-alternative-tools-fitdist-or-scipystatsfit10-scores)

    - [Q 4.1 What is the algorithm behind Fitdist (1 scores)](#q41-what-is-the-algorithm-behind-scipystatsfit-a-mle-or-b-map--1-score)

    - [Q 4.2 Adjust Fitdist for your problem (4 scores)](#q42-implement-scipystatsfit-for-estimating-which-is-the-explectation-of--4-scores)
    
    - [Q 4.3 Check your results (5 scores)](#q43-check-your-results--5-scores)

---

## Q1. Data generation (Total 6 scores)

**Scoring**

- **Data generation setting.**  Plan to generate the data .... 

- **Implementation.**  Data generation according to the above specifications, ie, the likelihood and priors (2 scores)

- **Data visualization.** Show that the distributions are right by histograms or any suitable visualization (4 scores)

---

###  Q1.1 Parameter settings ....  

Type down the python code to assign the values for 

- `K = 50` 
- `N = 2` 
- `M = 1000`

where  K is the number of random observation; N is the number of random variable dimension; M is the number of samples.

```python
K = 50
N = 2
M = 1000
```

The constants are
- `sigma_v` : $\sigma_V^2$ = 1
- `alpha` : $\alpha$ = 1
- `beta` : $\beta$ = 2
- `mu` : $\mu$ = a vector of zeros of size M x K x 1

```python
sigma_v = 4
alpha = 4
beta  = 2
mu    = np.zeros((M, 1, N)) + 0.5
```

---

###  Q1.2 Data generation according to the above specifications, ie, the likelihood and priors .... (2 scores)

Please follow the symbol belows to generate samples for $\bm{Y}$, $\bm{X}$, $N$, $\bm{\tau}$ which are RVs;  


- `tau` : $\bm{\tau}$ of size M x K x N 
- `X` : $\bm{X}$ of size M x K x N
- `Noise` : $N$      of size M x K x 1
- `Y` : $\bm{Y}$ of size M x K x N


For example .... 

``` 
Noise = randn(M, K, 1) 
X   = randn(M, K, N) 
Y   = X  + Noise
```
Note that above example is in Matlab and is incorrect ....  You are simulating everything in Python naka....

<!-- grade_id: q12_data_generation, points: 0, grade: false, solution: true -->
```python
# tau   = np.random.gamma ....
# X     = mu + ....
# Noise = np.random.randn ....
# Y     = X + Noise 


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q12_data_generation_testing, points: 2, grade: true, solution: false -->
```python
# Check that Y is an numpy array with the dimension M x K x N
assert  M == Y.shape[0]   
assert  K == Y.shape[1]  
assert  N == Y.shape[2]   

# Check that  X and tau are numpy arrays with the dimension M x 1 x N
assert tau.shape[0] == X.shape[0]  == M   
assert tau.shape[2] == X.shape[2]  == N  
assert tau.shape[1] == X.shape[1]  == 1 

# Check that Noise is the numpy array with the dimension M x K x 1  
assert Noise.shape[0] == M   
assert Noise.shape[1] == K   
assert Noise.shape[2] == 1
```

---

### Q1.3 Show that the distributions are correct by histograms or any suitable visualization (4 scores)

You can use the following histogram function to plot your data.

```python
def hist_plot(data_population_list, RV_label=["X"], color_list=['blue'], num_bins=100, dirpath=None, ploting_name=None): 
    ploting_title = "Histogram" 
    if ploting_name is None:
        ploting_name = ploting_title

    for sample_, RV_label_, color_ in zip(data_population_list, RV_label, color_list):
        plt.hist(sample_, bins=num_bins, density=False, color=color_, alpha = 0.4, label="%s [#%d]" % (RV_label_, len(sample_)))   

    plt.legend(loc='upper right')
    plt.ylabel("#Counts")
    plt.xlabel("Sample Values")
    plt.title(ploting_title) 
    plt.grid() 
    if dirpath is not None:
        plt.savefig(os.path.join(dirpath, "%s.png" % (ploting_name)))
        print("Save histrogram %s.png and .pdf at %s" % (ploting_name, dirpath))
    plt.show() 
```

```python
# Example usage... 
x_example_1 = np.random.randn(M, K, 1)
x_example_2 = np.random.randn(M, K, 1)

data_population_list = [x_example_1[:,0,0], x_example_2[:,0,0]]

RV_label_list = ["$X_1$",   "$X_2$"]
color_list    = ["red", "blue"]  

dirpath = os.getcwd()
ploting_name = "example_histrogram"
hist_plot(data_population_list, RV_label=RV_label_list, color_list=color_list, num_bins=100, dirpath=dirpath, ploting_name=ploting_name)
```

---

#### A. Show histrogram of `tau` .... Is it the right Gamma distribution? (1 score)

<!-- grade_id: q13_visualize_tau_hist, points: 0, grade: false, solution: true -->
```python
dirpath = os.getcwd()
ploting_name = "Q1.3.A.histrogram_of_tau_gamma_distribution"

RV_label_list = [r'$\tau_i$, $\alpha$=%d, $\beta$=%d' % (alpha, beta)]  ### < Note that you need to show the parameters
color_list    = ["purple"]  

# Your code starts here ...

# Hint!

# Define the data population as list of tau[:, 0, 0] < Note that you can use only the samples of a random sample in the first dimension.
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: q13_visualize_tau_hist_plot, points: 1, grade: true, solution: false -->
```python
hist_plot(data_population_list, RV_label=RV_label_list, color_list=color_list, num_bins=100, dirpath=dirpath, ploting_name=ploting_name)
```

---

#### B. Show histrogram of `X` .... Is it the right  Gaussian distribution? (1 scores)

<!-- grade_id: q13_visualize_X_hist, points: 0, grade: false, solution: true -->
```python
dirpath = os.getcwd()
ploting_name = "Q1.3.B.histrogram_of_X_gaussian_distribution"

RV_label_list = [r'$X_i$, $Gauss$ with $\mu$ and $\tau$']
color_list    = ["red"]   

# Your code starts here ...

# Hint!

# Define the data_population_list ...   look at the above hint

# YOUR CODE HERE
raise NotImplementedError()

# Your code ends here .

```

<!-- grade_id: q13_visualize_X_hist_plot, points: 1, grade: true, solution: false -->
```python
hist_plot(data_population_list, RV_label=RV_label_list, color_list=color_list, num_bins=100, dirpath=dirpath, ploting_name=ploting_name)
```

---

#### C. Show histrogram of `Noise` .... Is it the right  Gaussian distribution? (1 scores)

<!-- grade_id: q13_visualize_Noise_hist, points: 0, grade: false, solution: true -->
```python
dirpath = os.getcwd()
ploting_name = "Q1.3.C.histrogram_of_Noise_gaussian_distribution"

RV_label_list = [r'$N$, $Gauss$ with $0$ and $\sigma^2_v$']
color_list    = ["green"]   

# Your code starts here ...

# Hint!

# Define the data_population_list ...   

# YOUR CODE HERE
raise NotImplementedError()

# Your code ends here ..
```

<!-- grade_id: q13_visualize_Noise_hist_noise, points: 1, grade: true, solution: false -->
```python
hist_plot(data_population_list, RV_label=RV_label_list, color_list=color_list, num_bins=100, dirpath=dirpath, ploting_name=ploting_name)
```

---

#### D. Show that the histrogram of `Y` that is the summation of `X` and `N` ? (1 scores) 
- Plot the overlayed histrograms between `X` and `N`. You should use different colors for `X` and `N` ...
- Then, compare the histrogram of `Y` with the overlayed histrogram.

<!-- grade_id: q13_visualize_Y_hist, points: 0, grade: false, solution: true -->
```python
dirpath = os.getcwd()
ploting_name = "Q1.3.D.histrogram_of_Y_sum_of_two_distributions"

RV_label_list = ['$Y= X_i + N$', r'$X$, $Gauss$ with $\mu$ and $\tau$', r'$N$, $Gauss$ with $0$ and $\sigma^2_v$']
color_list    = ["blue", "red", "green"]

# Your code starts here ...

# Hint!
# data_population_list = list of Y, X, and Noise
# YOUR CODE HERE
raise NotImplementedError()


# Your code ends here ...
```

<!-- grade_id: q13_visualize_Y_hist_plot, points: 1, grade: true, solution: false -->
```python
hist_plot(data_population_list, RV_label=RV_label_list, color_list=color_list, num_bins=100, dirpath=dirpath, ploting_name=ploting_name)
```

---

## Q.2 Implementation of MAP/MLE Estimator (6 scores)

**Scoring**

- **Your analyses.** Analytical solutions for MLE and MAP for $X_i$....  (4 scores)

- **Your implementation.** Provide implementation for MLE and MAP for $X_i$ and then $\bm{x}$.... (8  scores)  

### Q.2.1 **Analytical Solutions of MAP/MLE Estimator for $X_i$**  (2 scores)

---

A. Analytical solutions for MLE for $X_i$ ? (1 scores)

<!-- grade_id: q21_mle, points: 1, grade: true, solution: true -->
$X_i = ...  

YOUR ANSWER HERE

B. Analytical solutions for MAP for $X_i$ (1 scores)

<!-- grade_id: q21_map, points: 1, grade: true, solution: true -->
$X_i = ... 

YOUR ANSWER HERE

---

### Q 2.2 **Implementation of MAP/MLE Estimator for $X_i$  and $x$**  (4 scores)

A. MLE implementation (1 scores)

<!-- grade_id: q22_implement_MLE, points: 0, grade: false, solution: true -->
```python
def x_MLE(Y, K):
    # YOUR CODE HERE
    raise NotImplementedError()
    return x_mle[:, np.newaxis,:]

```

B. MAP implementation (3 scores)

<!-- grade_id: q22_implement_MAP, points: 0, grade: false, solution: true -->
```python
def x_MAP(Y, K, mu, tau, sigma_v):
    # YOUR CODE HERE
    raise NotImplementedError()
    return x_map
```

---

## Q.3 Verify your results (8 scores)

---

### Q.3.1 Checking with respect to the generated $X$ (4 scores)

- Perform estimation for $X_i$ given $Y$ and the rest of parameters

<!-- grade_id: q31_run_MLE, points: 1, grade: true, solution: false -->
```python
x_mle =  x_MLE(Y, K)
```

<!-- grade_id: q31_run_MAP, points: 3, grade: true, solution: false -->
```python
x_map =  x_MAP(Y, K, mu, tau, sigma_v) 
```

- The histrogram comparisons between `x_mle`, `x_map`, and `X`. (4 scores)

<!-- grade_id: q31_hist_MLE_MAP, points: 0, grade: false, solution: true -->
```python
dirpath = os.getcwd()
ploting_name = "Q2.3 Histrogram_of_X_MLE_X_MAP_and_X"

RV_label_list = ['$X_i$',  "${X_i}_{MLE}$",  "${X_i}_{MAP}$"]
color_list    = ["blue",  "red",  "green"]

# Your code starts here ...

# Hint!
# data_population_list = list of X, x_mle, and x_map 
# You can use `hist_plot` to perform the task

# YOUR CODE HERE
raise NotImplementedError()

# Your code ends here ...
```

<!-- grade_id: q31_hist_MLE_MAP_plot, points: 4, grade: true, solution: false -->
```python
hist_plot(data_population_list, RV_label=RV_label_list, color_list=color_list, num_bins=100, dirpath=dirpath, ploting_name=ploting_name) 
```

---

### Q.3.2 Verify with $KL$-divergence ? (4 scores)

This section we will compare the KL divergence between `x_mle` and `X` versus `x_map` and `X`.

---

#### A. Before we go, can you provide the definition of KL-divergence ...(1 score) 

Let: 
- $P(x)$ is for the reference distribution.
- $Q(x)$ is for the testing distribution.

<!-- grade_id: q32_KL_divergence_def, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python
def KL_expo(P_samples, Q_samples ): 
    KL = np.sum(P_samples * np.log(P_samples / Q_samples))
    return KL
```

Prepare $P$ and $Q$ for KL divergence computation

```python
eps = 1e-10 # eps is used to up make sure that both P and Q will not be absolute zero at the same time, which can result in NaN.

# To calculate the KL-divergence empirically, we need to estimate the density for each distribution to be used as the input for KL-divergence calculation.

P_samples,_      = np.histogram(X[:,0,0],     bins=np.arange(-10,10,0.5), density=True)
Q_samples_mle ,_ = np.histogram(x_mle[:,0,0], bins=np.arange(-10,10,0.5), density=True) 
Q_samples_map,_  = np.histogram(x_map[:,0,0], bins=np.arange(-10,10,0.5), density=True) 

P_samples = P_samples + eps
Q_samples_mle = Q_samples_mle + eps 
Q_samples_map = Q_samples_map + eps 
```

---

#### B. What is the KL divergence  between X and X_mle ? (1 scores)

Hint!  KL_PQ_mle = KL_expo(P_samples, Q_samples_mle)

<!-- grade_id: q32_B_KL_X_Xmle, points: 0, grade: false, solution: true -->
```python
# KL_PQ_mle will be used to store the KL divergence

# Start your code here

# YOUR CODE HERE
raise NotImplementedError()

# your code ends here  
```

<!-- grade_id: q32_B_KL_X_Xmle_check, points: 1, grade: true, solution: false -->
```python
print("KL divergence between X and X_mle: %.10f" % KL_PQ_mle )

assert KL_PQ_mle > 0 and KL_PQ_mle <= 10
```

---

#### C. What is the KL divergence  between X and X_map ?  (2 scores)

<!-- grade_id: q32_B_KL_X_Xmap, points: 0, grade: false, solution: true -->
```python
# KL_PQ_map will be used to store the KL divergence

# Start your code here

# YOUR CODE HERE
raise NotImplementedError()

# your code ends here
```

<!-- grade_id: q32_B_KL_X_Xmap_checking, points: 2, grade: true, solution: false -->
```python
print("KL divergence between X and X_map: %.10f " % KL_PQ_map )
assert KL_PQ_map > 0 and KL_PQ_map <= 10 
```

---

## Q.4 Alternative tools (FitDist or `scipy.stats.fit`)....(10 scores)

---

### Q.4.1 What is the algorithm behind `scipy.stats.fit`...   (a) MLE or (b) MAP ?  (1 score)

Hint! You can check for the answer from the official website from scipy by clicking [here](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.fit.html)... 

<!-- grade_id: q4_scipy_stat_fit, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---

### Q.4.2 Implement `scipy.stats.fit` for estimating $X_i$ which is the explectation of $Y_i$. (4 scores)

Hint! You can use `stats.norm.fit(Y)` to check for the parameters of $Y$ as a Gaussian random samples. 

<!-- grade_id: q42_implement_scipy_stats, points: 0, grade: false, solution: true -->
```python
# your code starts here

import numpy as np
from scipy import stats 

mu_Y  = np.zeros_like(X)    # To store the estimated expectation of Y_i 
var_y = []                  # To store the estimated variance of Y_i

# YOUR CODE HERE
raise NotImplementedError()

# your code ends here   
```

<!-- grade_id: q42_implement_scipy_stats_check, points: 4, grade: true, solution: false -->
```python
# Check that  mu_Y and X are numpy arrays with the dimension M x 1 x N
assert mu_Y.shape[0] == X.shape[0]  == M   
assert mu_Y.shape[2] == X.shape[2]  == N  
assert mu_Y.shape[1] == X.shape[1]  == 1 
```

---

### Q.4.3 Check your results ... (5 scores)

#### A. Find KL divergence between reference $X_i$ and results from scipy.stats.fit (2 scores)

Hint ! Use KL_expo()

<!-- grade_id: q43_KL_X_vs_fitdist, points: 0, grade: false, solution: true -->
```python
# Start your code here 

# YOUR CODE HERE
raise NotImplementedError()

# end your code here 

```

<!-- grade_id: q43_KL_X_vs_fitdist_check, points: 1, grade: true, solution: false -->
```python
print("KL-Fitdist [%.10f]" % ( KL_PQ_fitdist))
```

<!-- grade_id: q43_KL_X_vs_fitdist_hidden, points: 1, grade: true, solution: false -->
```python

```

#### B. Find KL divergence between the estimator of your choice and one from scipy.stats.fit where you should compute both cases: (2 scores)

- Case 1.  P is the estimator of your choice (${X_{MLE}}_i$ or ${X_{MAP}}_i$), and Q is the results from scipy.stats.fit  (1 scores)

<!-- grade_id: q43_KL_Xc_vs_fitdist, points: 0, grade: false, solution: true -->
```python
# Start your code here 

# YOUR CODE HERE
raise NotImplementedError()

# end your code here 
```

<!-- grade_id: q43_KL_Xc_vs_fitdist_check, points: 1, grade: true, solution: false -->
```python
print("KL-P-CHOICE-Q-Fitdist [%.10f] " % (KL_P_CHOICE_Q_fitdist ))  
assert KL_P_CHOICE_Q_fitdist < 1e-3  
```

- Case 2.  switch P and Q. (1 scores)

<!-- grade_id: q43_KL_Xc_vs_fitdist_switch, points: 0, grade: false, solution: true -->
```python
# Start your code here  

# YOUR CODE HERE
raise NotImplementedError()
 
# end your code here 
```

<!-- grade_id: q43_KL_Xc_vs_fitdist_switch_check, points: 1, grade: true, solution: false -->
```python
print("KL-P-Fitdist-Q-CHOICE [%.10f] " % (KL_P_fitdist_Q_CHOICE)) 
assert KL_P_fitdist_Q_CHOICE < 1e-3  
```

#### C. Compare the histogram of the estimator of your choice and `scipy.stats.fit` (1 score)

<!-- grade_id: q43_hist_check, points: 0, grade: false, solution: true -->
```python
dirpath = os.getcwd()
ploting_name = "Q4.C.3 Histrogram_of_X_fitdist_vs_your_chosen_estimator"

RV_label_list = ['${X_i}_{fitdist}$',  "${X_i}_{choice}$" ]
color_list    = ["blue",  "red" ]

# Your code starts here ...

# Hint!
# data_population_list = list of mu_Y and estimator of your choice ( x_mle or x_map) 
# You can use `hist_plot` to perform the task

# YOUR CODE HERE
raise NotImplementedError()

# Your code ends here ...
```

<!-- grade_id: q43_hist_check_plot, points: 1, grade: true, solution: false -->
```python
hist_plot(data_population_list, RV_label=RV_label_list, color_list=color_list, num_bins=100, dirpath=dirpath, ploting_name=ploting_name) 
```